<a href="https://colab.research.google.com/github/kesarwani-shivam/Machine-Learning/blob/master/Linear_Model_using_LEAST_MEAN_SQUARE_and_MIN_NORM_solutions.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Function Approximation with a Linear Model
$y=f(x)=\sin x$ <br />
Approximate this function with a linear function of the form <br />
$y = w_0 + w_1 x + w_2 x^2$ <br />
in the range $x\in(0,2\pi)$

```python
import numpy as np
```

### Generate data

<!-- grade_id: cell-42de5e80af6ec099, points: 0, grade: false, solution: true -->
```python
def fx(x):
    '''
    Input:
        x: float (scalar) or np array
    Output:
        y: float (scalar) or np array; y=f(x)=sin(x)
        
    HINT: Use web search on how to take sine in numpy
    '''
    # YOUR CODE HERE
    y = []

    if type(x)==type(1) or type(x)==type(1.0):
      y = np.sin(x)
    else:
      for e in x:
        y.append(np.sin(e))
      y = np.array(y)
    return y
```

<!-- grade_id: cell-f27e4f64b287cd94, points: 1, grade: true, solution: false -->
```python
'''Testing'''
def test_fx():
    assert fx(0)==0
    assert np.isclose(fx(np.pi/2), 1, atol=0.001)
    assert np.isclose(fx(np.pi/4), 0.707, atol=0.001)
    print('Test passed', '\U0001F44D')
test_fx()
```

Output:
```
Test passed 👍
```

<!-- grade_id: cell-284fa421c7c60677, points: 0, grade: false, solution: true -->
```python
def generateX_fx(N):
    '''
    Generate N random points in the range (0,1)
    Input:
        N: int
    Outputs:
        x: np vector of shape (N,)
        y: np vector of shape (N,); y=f(x)
    '''
    # YOUR CODE HERE
    x = np.random.rand(N)
    yd = fx(x)
    return x, yd
```

<!-- grade_id: cell-70c27081df964f24, points: 2, grade: true, solution: false -->
```python
'''Testing'''
def test_generateX_fx():
    x, yd = generateX_fx(20)
    assert x.shape[0]==20
    assert yd.shape[0]==20
    print('Test passed', '\U0001F44D')
test_generateX_fx()
```

Output:
```
Test passed 👍
```

### Plotting

```python
def plotdata():
    x, yd = generateX_fx(20)
    import matplotlib.pyplot as plt
    i = x.argsort()
    plt.plot(x[i],yd[i])
    plt.ylabel('f(x)')
    plt.xlabel('x')
# plotdata()
```

## Training
We will now define a linear model to estimate the above function, and train it with all the data

### Define model
$y = \sum_{d=0}^D w_d x^d$, here, $D$ is the degree of the model, say D=2 <br />

In matrix representation: y = Aw <br/>

<!-- grade_id: cell-61478073f3568175, points: 0, grade: false, solution: true -->
```python
def createA(x, D):
    '''
    Create the matrix A with degree D
    Input:
        x: np vector of shape (N,)
        D: degree of the model
    Output:
        A: np matrix of shape (N,D+1)
    '''
    # YOUR CODE HERE
    A= []
    for i in range(len(x)):
      a = []
      for j in range(0,D+1):
        a += [x[i]**j]
     
      A += [a]
    A = np.array(A)
    
    
    return A
```

<!-- grade_id: cell-c1a05c3a2dc0aa06, points: 4, grade: true, solution: false -->
```python
'''Testing'''
def test_createA():
    x = np.array([1,2,3])
    A = createA(x, 2)
    assert A.shape==(3,3)
    assert np.all(A[0,:]==1)
    assert np.all(A==np.array([[1.0, 1.0, 1.0], [1.0, 2.0, 4.0], [1.0, 3.0, 9.0]]))
    print('Test passed', '\U0001F44D')
test_createA()
```

Output:
```
Test passed 👍
```

### Estimate Weights
Estimate w from yd and A. Use the least square solution and min-norm solution you learnt in the video

<!-- grade_id: cell-9bdf45f987ecf410, points: 0, grade: false, solution: true -->
```python
def train_w(A, yd):
    '''
    Inputs:
        A: np matrix of shape (N,D+1)
        yd: np vector of shape (N,)
    Output:
        w: np vector of shape (D+1,)
    '''
    # YOUR CODE HERE
    N, D1 = A.shape
    if N>D1-1:
      w = np.linalg.inv(A.T.dot(A)).dot(A.T.dot(yd))
    elif N<D1-1:
      w = A.T.dot(np.linalg.inv(A.dot(A.T)).dot(yd))
    
    
    return w


```

<!-- grade_id: cell-4fd1ee72464aa0d8, points: 4, grade: true, solution: false -->
```python
'''Testing'''
def test_train_w():
    A = np.array([[1.0, 1.0, 1.0], [1.0, 2.0, 4.0], [1.0, 3.0, 9.0]])
    yd = np.array([1,2,3])
    assert np.all(np.isclose(train_w(A, yd), np.array([0,1,0]), atol=.001))
    yd = yd*yd
    assert np.all(np.isclose(train_w(A, yd), np.array([0,0,1]), atol=.001))
    print('Test passed', '\U0001F44D')
test_train_w()
```

Output:
```
Test passed 👍
```

### Estimate y from the model, given x

<!-- grade_id: cell-18a403e9e2d75c54, points: 0, grade: false, solution: true -->
```python
def predict_y(w, x):
    '''
    Inputs:
        w: np vector of shape (D+1,)
        x: np vector of shape (N,)
    Outputs:
        y: np vector of shape (N,); y=Aw
    '''
    # YOUR CODE HERE
    A = createA(x,2)
    y = A.dot(w)
    return y
```

<!-- grade_id: cell-66e9156fdbab1ffe, points: 3, grade: true, solution: false -->
```python
def test_predict_y():
    w = np.array([1,2,3])
    x = np.array([0.2,0.5,0.6])
    y = predict_y(w, x)
    assert np.all(np.isclose(y, np.array([1.52, 2.75, 3.28]), atol=0.001))
    print('Test passed', '\U0001F44D')
test_predict_y()
```

Output:
```
Test passed 👍
```

### Estimate Error
Find E as the mean squared error

<!-- grade_id: cell-430f96298f1ea137, points: 0, grade: false, solution: true -->
```python
def compute_mse(y, yd):
    '''
    Inputs:
        y: np vector of shape (N,); y=Aw
        yd: np vector of shape (N,); yd=f(x), ie., desired or true value
    Output:
        mse: mean squared error
    '''
    # YOUR CODE HERE
    mse = sum([(y[i]-yd[i])**2 for i in range(len(y))])
    mse = mse/len(yd)
    return mse
```

<!-- grade_id: cell-2ce30fda249b5070, points: 2, grade: true, solution: false -->
```python
'''Testing'''
def test_compute_mse():
    yd = np.array([1,0,1,0])
    y = np.array([.8,.2,.8,.2])
    assert np.isclose(compute_mse(y, yd), 0.04, atol=0.003)
    print('Test passed', '\U0001F44D')
test_compute_mse()
```

Output:
```
Test passed 👍
```

### Plotting


```python
def plotModel(x, y, yd):
    import matplotlib.pyplot as plt
    i = x.argsort()
    plt.figure()
    plt.plot(x[i],y[i],'g-o')
    plt.plot(x[i],yd[i],'r-o')
    plt.ylabel('f(x)')
    plt.xlabel('x')
    plt.legend(['estimated', 'true'])
```

## Could you test it on new (test) points?
Using the above functions:
- Generate 30 test points
- Estimate y using the linear model w you found above. Do NOT train again
- Compare y with yd to find the mean-squared error <br />

<!-- grade_id: cell-4a89399915bf1752, points: 0, grade: false, solution: true -->
```python
def testModel(Ntest, w):
    '''
    Inputs:
        Ntest: number of test samples to be generated
        w: np array of size (D+1,)
    Outputs:
        x: np array of size (N,)
        y: np array of size (N,)
        yd: np array of size (N,)
        mse: scalar float
    '''
    # YOUR CODE HERE
    x, yd = generateX_fx(Ntest)
    y = predict_y(w, x)
    mse = compute_mse(y,yd)
    return x, y, yd, mse
```

<!-- grade_id: cell-7b7b3892f3df2282, points: 4, grade: true, solution: false -->
```python
'''Testing'''
def test_testModel():
    N = 30
    D = 10
    x, y, yd, w, mse = trainModel(N,D)
    assert x.shape[0]==N
    assert w.shape[0]==D+1
    print(w)
    print('MSE on train data = ', mse)
    plotModel(x,y,yd)
    Ntest = 50
    x, y, yd, mse = testModel(Ntest,w)
    print('MSE on test data = ', mse)
    plotModel(x,y,yd)
    print('Test passed', '\U0001F44D')
test_testModel()
```

Output:
```
<ipython-input-61-86aca7545922> in <module>()
     14     plotModel(x,y,yd)
     15     print('Test passed', '\U0001F44D')
---> 16 test_testModel()

<ipython-input-61-86aca7545922> in test_testModel()
      3     N = 30
      4     D = 10
----> 5     x, y, yd, w, mse = trainModel(N,D)
      6     assert x.shape[0]==N
      7     assert w.shape[0]==D+1

NameError: name 'trainModel' is not defined
```

# ADVANCED(Optional)

## Does the performance improve with increasing the number of training points?
- Repeat the training with different no. of samples. Take N=3, 5, 10, 30; and D=5
- Plot training error vs N
- Plot test error vs N

```python

```

## Does the performance improve with increasing the degree D of the model?
- Repeat the training with different degree models. Take D=0, 1, 2, 4, 7, 12, and N=10
- Plot training error vs D
- Plot test error vs D
- Plot y vs x for each D

```python

```